```python
import csv
import json
import numpy as np
import warnings
```

### Exercise 1 (4 points)

This exercise is a quiz with 10 multiple answer questions.
In the cell below, we've declared a dictionary called `answers`.

You should fill in that dictionary with your answers, using as keys the question number, like `question_x`, and as values the number from `1` to `4` that corresponds to the right answer.
For each question, you should provide only one answer, i.e, the dict values should have type **int**.

For example, if you want to answer Question 1 with choice number 2, then you do:
```
answers["question_1"] = 2
```

Each correct answer is worth 0.4 points.

```python
answers = {}
```

#### Question 1
Which command can you use to print the path of the current working directory in the terminal of a Unix based machine?
1. cd .
2. ls .
3. pwd
4. cd

<!-- grade_id: cell-c87b80092d254365, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
answers["question_1"] = 3
### END SOLUTION


# answers["question_1"] = ...
```

#### Question 2
If you want to remove a file `data.csv` located in a sub directory `subdir` using the terminal of a Unix based machine, which of these commands can you use?
1. rm data.csv
2. rm subdir/data.csv
3. rm -r subdir
4. mv data.csv subdir

<!-- grade_id: cell-c1fc749530d189a6, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
answers["question_2"] = 2
### END SOLUTION


# answers["question_2"] = ...
```

#### Question 3
In the terminal of a Unix based machine, which of these commands would you use to print the initial 5 lines of a file named `data.csv`?
1. head -5 data.csv
2. tail -5 data.csv
3. cat data.csv
4. ls -la data.csv

<!-- grade_id: cell-98b6eeed9f5eb544, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
answers["question_3"] = 1
### END SOLUTION


# answers["question_3"] = ...
```

#### Question 4
In the terminal of a Unix based machine, if you want to see the size of the files in the current directory, in a human readable format, which command is the best option?
1. ls
2. ls .
3. ls -la
4. ls -lh

<!-- grade_id: cell-151e77358095531d, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
answers["question_4"] = 4
### END SOLUTION


# answers["question_4"] = ...
```

#### Question 5
If you are working on a project using git as a version control system and you changed a file `utils.py`, how do you commit your changes to a local branch?
1. `git add utils.py & git commit -m "updates helper function"`
2. `git status & git commit -m "updates helper function"`
3. `git status & git push`
4. `git add utils.py & git push`

<!-- grade_id: cell-fc82bf7c933cc552, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
answers["question_5"] = 1
### END SOLUTION


# answers["question_5"] = ...
```

#### Question 6
If you are using git to version control a project and you want to avoid that files with names starting with `data` are commited, what do you do?
1. Run the command `git rm data`
2. Add `data*` to the `.gitignore` file
3. Run the command `git rm data*`
4. Add `data` to the `.gitignore` file

<!-- grade_id: cell-28d4124915005816, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
answers["question_6"] = 2
### END SOLUTION


# answers["question_6"] = ...
```

#### Question 7
Which one of the following sentences is false?
1. Dictionaries can have duplicated keys.
2. Dictionaries can have duplicated values.
2. Lists can have duplicated elements.
3. Tuples can have duplicated elements.

<!-- grade_id: cell-8bf8a029ce326cc8, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
answers["question_7"] = 1
### END SOLUTION


# answers["question_7"] = ...
```

#### Question 8
To interrupt a `for` loop iteration and start the next iteration, which instruction would you use?
1. stop
2. continue
3. break
4. start

<!-- grade_id: cell-fe303ff800b35243, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
answers["question_8"] = 2
### END SOLUTION


# answers["question_8"] = ...
```

#### Question 9
How do you import function `my_function` from module `my_module`?
1. import my_function
2. import my_module
3. from my_function import my_module
4. from my_module import my_function

<!-- grade_id: cell-c2eef7f0328b8ba7, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
answers["question_9"] = 4
### END SOLUTION


# answers["question_9"] = ...
```

#### Question 10
In which occasion should you use a try/except block and raise an exception?
1. to run a block of code based on a condition
2. to avoid that a program fails
3. to run a block of code repeatedly
4. to stop a program if an unexpected condition is found

<!-- grade_id: cell-e55806749e917025, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
answers["question_10"] = 4
### END SOLUTION


# answers["question_10"] = ...
```

<!-- grade_id: cell-e8e8abee625de5d7, points: 4, grade: true, solution: false -->
```python
# This cell runs some hidden tests.
# To see how much you scored in this question, please submit your notebook in the portal.

### BEGIN HIDDEN TESTS
from hidden_tests import test_exercise_1
test_exercise_1(answers)
### END HIDDEN TESTS
```

### Exercise 2

Consider the csv file `products.csv`, included in the zip file you downloaded with this notebook. This file contains four columns describing products available at Supermarket X:
- `name`: product name
- `category`: product category
- `price_kg`: price per kilogram of product, in units of €/kg
- `kcal_kg`: nutritional energy per kilogram of product, in units of kcal/kg

#### Part I (2 points)

Implement a function that reads a file with the same format as the `products.csv` file, and stores the data in a dictionary with the following structure:

```
{
    "[product_name]": {
        "idx": row number in the input file, not considering the header,
               i.e, the first row with data values has idx=0 (type: int),
        "category": category (type: str)
        "price_kg": price per kg value (type: float),
        "kcal_kg": nutritional energy per kg value (type: float)
    },
}

```

For example, if the first line in `products.csv` relates to product with name `"Banana"`, category `"fruit"`, price per kg `1.1` and nutritional energy per kg `890`, then the dictionary should contain the item

```
{
    "Banana": {
        "idx": 0,
        "category": "fruit"
        "price_kg": 1.1,
        "kcal_kg": 890.0,
    },
}

```

The function should be called `read_products`. It should receive an argument `filepath`, that indicates the path to the file that the function should read the data from, and return the dictionary with product data.

<!-- grade_id: cell-fd56150b3fe08280, points: 0, grade: false, solution: true -->
```python
def read_products(filepath):
    """
    Reads the file in filepath, parses it and returns the data in a dictionary.

    Parameters:
    filepath (str): Path to the input file to be parsed.

    Returns:
    products_parsed_data (dict): Products dataset stored as a dictionary, where
    the keys are product names and the values are data about the product.
    """
    
    ### BEGIN SOLUTION
    import csv

    products_parsed_data = {}
    with open(filepath, "r") as fp:
        reader = csv.reader(fp)
        for i, row in enumerate(reader):
            if i == 0:
                continue

            products_parsed_data[row[0]] = {
                "idx": i - 1,
                "category": row[1],
                "price_kg": float(row[2]),
                "kcal_kg": float(row[3]),
            }
    
    return products_parsed_data

    ### END SOLUTION
```

<!-- grade_id: cell-3169e339199d3588, points: 2, grade: true, solution: false -->
```python
# This cell has tests to evaluate your answer
# If the cell runs successfully, and no errors are raised, then you get the full score for this exercise
# Otherwise, you are not granted any points


products_data = read_products("products.csv")

assert len(products_data) == 13

assert "Avocado" in products_data
avocado = products_data["Avocado"]
assert avocado["idx"] == 6
assert avocado["category"] == "fruit"
assert isinstance(avocado["price_kg"], float)
assert isinstance(avocado["kcal_kg"], float)
np.testing.assert_almost_equal(avocado["price_kg"], 4.5, decimal=1)
np.testing.assert_almost_equal(avocado["kcal_kg"], 1600, decimal=1)

from more_tests import test_exercise_2_I
test_exercise_2_I(read_products)
```

#### Part II (3 points)

Consider now that we want to prepare a recipe that uses the products parsed from the file `products.csv`. Consider that this recipe is represented in a dictionary, where keys are product names and values are the weights (in kg) of the corresponding products required to prepare the recipe.

Implement a function called `recipe_price_kcal` that calculates the total price and nutritional energy of the recipe. The function should receive as input two dictionaries, representing:

i) the products, using the same structure as the dictionary we created in Part I of this exercise, and

ii) the recipe (in which keys are product names and values are the amount required, in kg).

You can assume that all products required to prepare the recipe are available in the products dictionary.
The function should return the price and nutritional energy of the recipe in a tuple.

<!-- grade_id: cell-1346db6f517997b7, points: 0, grade: false, solution: true -->
```python
def recipe_price_kcal(products, recipe):
    """
    Computes the total price and nutritional energy of a recipe.

    Parameters:
    products (dict): Dictionary with the products dataset.
    recipe (dict): Dictionary with recipe data.

    Returns:
    price_kcal (tuple): Tuple with the total price and nutritional energy of 
    the recipe in the first and second elements, respectively.
    """
    
    ### BEGIN SOLUTION
    price = 0
    kcal = 0
    
    for product, quantity in recipe.items():
        price += products[product]["price_kg"] * quantity
        kcal += products[product]["kcal_kg"] * quantity
    
    return (price, kcal)
    ### END SOLUTION
```

<!-- grade_id: cell-5493c05da85f0a2d, points: 3, grade: true, solution: false -->
```python
# This cell has tests to evaluate your answer
# If the cell runs successfully, and no errors are raised, then you get the full score for this exercise
# Otherwise, you are not granted any points


iogurt_cake_products_data = {
    "Flour": {"idx": 0, "category": "cereal", "price_kg": 0.5, "kcal_kg": 3500.0},
    "Sugar": {"idx": 1, "category": "other", "price_kg": 0.7, "kcal_kg": 4000.0},
    "Egg": {"idx": 2, "category": "protein", "price_kg": 6.5, "kcal_kg": 1320.0},
    "Iogurt": {"idx": 3, "category": "dairy", "price_kg": 2.3, "kcal_kg": 1300.0},
    "Sunflower oil": {"idx": 4, "category": "oil", "price_kg": 1.7, "kcal_kg": 8300.0},
}

iogurt_cake_recipe = {
    "Flour": 0.3,
    "Sugar": 0.3,
    "Egg": 0.3,
    "Iogurt": 0.2,
    "Sunflower oil": 0.1
}

price_kcal = recipe_price_kcal(iogurt_cake_products_data, iogurt_cake_recipe)

assert isinstance(price_kcal, tuple)
assert len(price_kcal) == 2
np.testing.assert_almost_equal(price_kcal[0], 2.94, decimal=2)
np.testing.assert_almost_equal(price_kcal[1], 3736, decimal=2)

from more_tests import test_exercise_2_II
test_exercise_2_II(recipe_price_kcal)
```

### Exercise 3

In this exercise, we will use object oriented programming concepts to model a **bus** and determine its ocupation along a given route.

#### Part I  (4 points)

Create a `Bus` class. It should store the following information about the bus:
- `idx`: the bus identification index
- `dimensions`: the bus width, length and height, in meters
- `max_capacity`: the maximum number of passengers of the bus
- `num_passengers`: the current number of passengers of the bus

In creating the `Bus` class, consider that:
- the `idx`, `max_capacity` and `num_passengers` attributes should be represented with non-negative integers.
- the `dimensions` attribute should be a list of floats that indicates the width, length and height of the bus in this order; for example, if the bus is 2.5 m wide, 15 m long and 3 m tall, the list should be `[2.5, 15, 3]`.
- when a `Bus` instance is created, its `idx`, `dimensions` and `max_capacity` should be specified; however, `num_passengers` should always be initially set to 0.

The `Bus` class should have only one method, `update_passengers`. This method updates the number of passengers of the bus when it stops. It should receive two keyword arguments:
- `leave`: the number of passengers that leaves the bus at the stop; defaults to zero
- `enter`: the number of passengers to enter the bus at the stop; defaults to zero

This method should only update the `num_passengers` of the `Bus` instance, and thus does not return anything. In updating `num_passengers`, remember that the maximum number of passengers of the bus should not exceed `max_capacity`, even if that means leaving a fraction of the `enter` passengers waiting for the next bus. If this happens, a warning should be issued with the message "X people were left waiting for the next bus", where X is the number of people that could not enter the bus at the stop. To issue this warning, you should use the [**warnings.warn** method](https://docs.python.org/3/library/warnings.html#warnings.warn) (note that we have already imported the `warnings`module for you).

<!-- grade_id: cell-e4bd0792c4fcba1b, points: 0, grade: false, solution: true -->
```python
class Bus:
    ### BEGIN SOLUTION
    def __init__(self, idx, dimensions, max_capacity):
        self.idx = idx
        self.dimensions = dimensions
        self.max_capacity = max_capacity
        self.num_passengers = 0
        
    def update_passengers(self, leave=0, enter=0):
        d_passengers = self.num_passengers - leave + enter - self.max_capacity
        if d_passengers > 0:
            warnings.warn(f"{d_passengers} people were left waiting for the next bus")
        self.num_passengers -= leave
        self.num_passengers = min(self.num_passengers + enter, self.max_capacity)
    ### END SOLUTION
```

<!-- grade_id: cell-f0f762d780f9cb8b, points: 4, grade: true, solution: false -->
```python
# This cell has tests to evaluate your answer
# The output of this cell will show how many points you scored and,
# in case you don't have the correct answer, some error messages


bus = Bus(123, [2.5, 15, 3], 30)

score = 0

try:
    assert isinstance(bus, Bus), "Bus type is wrong!"
    assert bus.idx == 123, "Idx is wrong!"
    assert bus.dimensions == [2.5, 15, 3], "Dimensions are wrong!"
    assert bus.max_capacity == 30, "Max capacity is wrong!"
    assert bus.num_passengers == 0, "Number of passengers is wrong!"
except AssertionError as e:
    print(f"Failed first test: {e}")
else:
    score += 1.5
    
try:
    with warnings.catch_warnings(record=True) as w:
        bus.update_passengers(leave=0, enter=20)
        assert bus.num_passengers == 20, "Number of passengers after first stop is wrong!"

        bus = Bus(123, [2.5, 15, 3], 30)
        bus.num_passengers = 20
        bus.update_passengers()
        assert bus.num_passengers == 20, "Number of passengers after second stop is wrong!"

        bus = Bus(123, [2.5, 15, 3], 30)
        bus.num_passengers = 20
        bus.update_passengers(leave=5)
        assert bus.num_passengers == 15, "Number of passengers after third stop is wrong!"

except AssertionError as e:
    print(f"Failed second test: {e}")
else:
    score += 1.5

try:    
    with warnings.catch_warnings(record=True) as w:
        bus = Bus(123, [2.5, 15, 3], 30)
        bus.num_passengers = 20
        bus.update_passengers(leave=5)
        assert bus.num_passengers == 15, "Number of passengers after fourth stop is wrong!"
        assert len(w) == 0, "Warning is incorrectly raised, no passengers are waiting!"
        
        bus = Bus(123, [2.5, 15, 3], 30)
        bus.num_passengers = 15
        bus.update_passengers(enter=20)
        assert bus.num_passengers == 30, "Number of passengers after fifth stop is wrong!"
        assert len(w) == 1, "A warning was not raised after fifth stop!"
        assert str(w[-1].message) == "5 people were left waiting for the next bus", "The message of the warning raised after fifth stop is wrong!"
        
    
except AssertionError as e:
    print(f"Failed third test: {e}")
else:
    score += 1

print(f"Your score is {score}.")

score
```

#### Part II (2 points)

We will now track the number of passengers in a bus along a route.

Write a function called `track_passengers` that receives two arguments:
- `bus`: an instance of `Bus`
- `stops`: a list of dictionaries with keys `leave` and `enter` that indicate the number of passengers to leave and enter the bus at each stop; for example `[{"enter": 10, "leave": 0}, {"enter": 5, "leave": 2}]` represents a sequence of two stops: 10 passengers enter the bus in the first stop, whereas in the second 2 passengers leave and 5 enter the bus. 

The function should return a list `passengers`, where the elements are the number of passengers after each bus stop in `stops`. The list should thus have the same length as `stops`.

<!-- grade_id: cell-4cbf77ea8554d92f, points: 0, grade: false, solution: true -->
```python
def track_passengers(bus, stops):
    """
    Tracks the number of passengers in a bus along a route.

    Parameters:
    bus (Bus): Bus to be tracked.
    stops (list): List of dictionaries representing the stops (i.e. the number
    of passengers to hop in/out of the bus) of the route.

    Returns:
    passengers (list): Number of passengers after each stop in the bus route.
    """
    ### BEGIN SOLUTION
    passengers = []
    
    for stop in stops:
        bus.update_passengers(stop["leave"], stop["enter"])
        passengers.append(bus.num_passengers)

    return passengers
    ### END SOLUTION
```

<!-- grade_id: cell-1dac271e0fc5c54c, points: 2, grade: true, solution: false -->
```python
# This cell has tests to evaluate your answer
# If the cell runs successfully, and no errors are raised, then you get the full score for this exercise
# Otherwise, you are not granted any points


bus = Bus(123, [2.5, 15, 3], 20)

stops = [
    {"enter": 10, "leave": 0},
    {"enter": 2, "leave": 0},
    {"enter": 5, "leave": 8},
    {"enter": 15, "leave": 2},
    {"enter": 0, "leave": 5},
]

with warnings.catch_warnings():
    warnings.simplefilter("ignore")
    passengers = track_passengers(bus, stops)
    
    assert isinstance(passengers, list)
    assert len(passengers) == 5
    assert passengers == [10, 12, 9, 20, 15]
    
    from more_tests import test_exercise_3_II
    test_exercise_3_II(track_passengers, Bus)
```

#### Part III (1 point)

You have probably heard about [The Magic School Bus](https://en.wikipedia.org/wiki/The_Magic_School_Bus_(TV_series)). This was a special school bus in which Miss Frizzle went on exciting field trips with her class, discovering "locations, creatures, time periods and more \[...\] about the wonders of science along the way".

One of the special perks of The Magic School Bus is that it was able to shrink and enter the human body, allowng Miss Frizzle and the students to go on a field trip to learn about cells! How exciting is that?

<img src="./images/magic_school_bus.jpg" style="width:400px;">

In this exercise, we will write a special class for The Magic School Bus, `MagicSchoolBus`, that inherits the properties of our `Bus` class. `MagicSchoolBus` should inherit all attributes and methods of the `Bus` class, and implement one extra method, called `shrink`, that shrinks the dimensions of the bus. This method takes as argument `factor`, a float representing the shrinking factor.

Shrinking a bus by a factor of 10 means that all its dimensions are reduced to a tenth of their original value. For example, shrinking a bus of dimensions `[2.5, 15, 3]` by a factor of 10 updates the `dimensions` atributes to `[0.25, 1.5, 0.3]`.

<!-- grade_id: cell-8d7bbbf2871cb57a, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
class MagicSchoolBus(Bus):
    def shrink(self, factor):
        for i in range(len(self.dimensions)):
            self.dimensions[i] /= factor
### END SOLUTION
```

<!-- grade_id: cell-0ad7e152e2d161e8, points: 1, grade: true, solution: false -->
```python
# This cell has tests to evaluate your answer
# The output of this cell will show how many points you scored and,
# in case you don't have the correct answer, some error messages


bus = MagicSchoolBus(123, [2.5, 15, 3], 30)

score = 0

try:
    assert type(bus) == MagicSchoolBus, "Magic school bus type is wrong!"
    assert isinstance(bus, Bus), "Magic school bus parent type is wrong!"
    
    assert bus.idx == 123, "Idx is wrong!"
    assert bus.dimensions == [2.5, 15, 3], "Dimensions are wrong!"
    assert bus.max_capacity == 30, "Max capacity is wrong!"
    assert bus.num_passengers == 0, "Number of passengers is wrong!"
except AssertionError as e:
    print(f"Failed first test: {e}")
else:
    score += 0.5
    
try:
    bus.shrink(factor=5)
    np.testing.assert_almost_equal(bus.dimensions, [x / 5 for x in [2.5, 15, 3]], decimal=2)
except AssertionError:
    print("Failed second test: Magic school bus dimensions are wrong!")
else:
    score += 0.5
    
print(f"Your score is {score}.")

score
```

### Exercise 4

[Minesweeper](https://en.wikipedia.org/wiki/Minesweeper_(video_game)) is a puzzle video game which goal is to clear a board containing hidden mines without detonating any of them, with help from clues about the number of neighboring mines in each field.

<img src="./images/minesweeper.png" style="width:400px;">

In this exercise, we will create the map of clues from the map of mines. The map of mines is given in a matrix that represents the board. In this matrix, cells with mines are filled the string `"*"`, and cells without mines are filled with the string `"-"`. 

Here is an example of a map of mines:

```python
mines_map = np.array([
    ["-", "*", "-", "-"],
    ["-", "-", "-", "-"],
    ["-", "-", "*", "-"],
    ["*", "-", "-", "-"],
    ["-", "-", "-", "-"],
])
```

Each cell of the map of clues should have the number of mines in neighbouring cells of the map of mines. For the map of mines above, the map of clues should be:

```python
clues_map = np.array([
    [ 1, -1,  1,  0],
    [ 1,  2,  2,  1],
    [ 1,  2, -1,  1],
    [-1,  2,  1,  1],
    [ 1,  1,  0,  0],
])
```

Note that: 
- all cells with mines have been marked in the map of clues with -1.
- matrix boundaries are treated as closed (i.e. `clues_map[4, 1]` does not know about the mine in `mines_map[0, 1]`)

#### Part I (2 points)

When navigating a map of mines, it's useful to know which are the cells adjacent to the current one.

For example, consider the map of mines presented above, and the mine in position `(0,1)`. The cells adjacent to the mine are the ones flagged with an `O`:

```python
mines_map = np.array([
    ["O", "*", "O", "-"],
    ["O", "O", "O", "-"],
    ["-", "-", "*", "-"],
    ["*", "-", "-", "-"],
    ["-", "-", "-", "-"],
])
```

This means that, the positions of the cells adjacent to the mine are the following ones:
- `(0,0)`
- `(1,0)`
- `(1,1)`
- `(1,2)`
- `(0,2)`

For this exercise, write a function named `get_adjacent_cells`, that receives two arguments:
- `cell_position`, represented as a tuple where the first element is the row and the second is the column, and
- `mines_map_shape`, represented as a tuple where the first element is the number of rows in the mines map and the second is the number of columns

The function should return a list with the positions of the adjacent cells, where each position is represented as a tuple where the first element is the row and the second is the column.

You can assume that the map of mines is not empty, i.e, it has at least one row and one column, and that the `cell_position` is inside the map.

<!-- grade_id: cell-e59124525f4d1a44, points: 0, grade: false, solution: true -->
```python
def get_adjacent_cells(cell_position, mines_map_shape):
    """
    Determines the positions of the cells adjacent to a given cell in a map of mines.

    Parameters:
    cell_position (tuple): Position of a cell.
    mines_map_shape (tuple): Shape of the mines map.

    Returns:
    adjacent_cells_positions (list): List with the positions of the adjacent cells.
    """

    ### BEGIN SOLUTION
    import itertools

    x, y = cell_position
    
    xs = [i for i in [x-1, x, x+1] if 0 <= i < mines_map_shape[0]]
    ys = [i for i in [y-1, y, y+1] if 0 <= i < mines_map_shape[1]]

    adjacent_cells = itertools.product(xs, ys)
    return [cell for cell in adjacent_cells if cell != cell_position]
    ### END SOLUTION
```

<!-- grade_id: cell-4f35e1ec6ac365eb, points: 2, grade: true, solution: false -->
```python
# This cell has tests to evaluate your answer
# If the cell runs successfully, and no errors are raised, then you get the full score for this exercise
# Otherwise, you are not granted any points

cell_position = (0, 0)
mines_map_shape = (5, 4)
answer = get_adjacent_cells(cell_position, mines_map_shape)
expected_answer = [(0, 1), (1, 0), (1, 1)]
assert isinstance(answer, list)
for position in answer:
    assert isinstance(position, tuple)
assert set(answer) == set(expected_answer)

cell_position = (0, 0)
mines_map_shape = (1, 1)
answer = get_adjacent_cells(cell_position, mines_map_shape)
expected_answer = []
assert isinstance(answer, list)
assert set(answer) == set(expected_answer)

cell_position = (1, 1)
mines_map_shape = (3, 3)
answer = get_adjacent_cells(cell_position, mines_map_shape)
expected_answer = [(0, 0), (0, 1), (0, 2), (1, 0), (1, 2), (2, 0), (2, 1), (2, 2)]
assert isinstance(answer, list)
for position in answer:
    assert isinstance(position, tuple)
assert set(answer) == set(expected_answer)

from more_tests import test_exercise_4_I
test_exercise_4_I(get_adjacent_cells)
```

#### Part II (1 point)

Write a function, named `get_clues_map` that receives as argument a numpy matrix `mines_map` that represents the map of mines and returns the corresponding map of clues `clues_map`.

You can assume that `mines_map` is not empty, i.e, it has at least one row and one column.

**Hint**: you may (and should!) use the function you wrote in Part I.

<!-- grade_id: cell-c25d9a32062cfeb0, points: 0, grade: false, solution: true -->
```python
def get_clues_map(mines_map):
    """
    Determines the map of clues from the map of mines.

    Parameters:
    mines_map (np.ndarray): Matrix with map of mines.

    Returns:
    clues_map (np.ndarray): Matrix with map of clues.
    """
    ### BEGIN SOLUTION
    clues_map = np.zeros(mines_map.shape)
    
    for i in range(mines_map.shape[0]):
        for j in range(mines_map.shape[1]):
            if mines_map[i][j] == "*":
                clues_map[i][j] = -1
                continue

            mines_count = 0
            for ii, jj in get_adjacent_cells((i, j), mines_map.shape):
                if mines_map[ii][jj] == "*":
                    mines_count += 1

            clues_map[i][j] = mines_count

    return clues_map
    ### END SOLUTION
```

<!-- grade_id: cell-6993496e56ae5877, points: 1, grade: true, solution: false -->
```python
# This cell has tests to evaluate your answer
# If the cell runs successfully, and no errors are raised, then you get the full score for this exercise
# Otherwise, you are not granted any points


mines_map = np.array([
    ["-", "*", "-", "-"],
    ["-", "-", "-", "-"],
    ["-", "-", "*", "-"],
    ["*", "-", "-", "-"],
    ["-", "-", "-", "-"],
])

clues_map = get_clues_map(mines_map)

assert isinstance(clues_map, np.ndarray)
assert clues_map.shape == mines_map.shape
np.testing.assert_array_equal(
    clues_map, 
    np.array([
        [ 1, -1,  1,  0],
        [ 1,  2,  2,  1],
        [ 1,  2, -1,  1],
        [-1,  2,  1,  1],
        [ 1,  1,  0,  0],
    ])
)

from more_tests import test_exercise_4_II
test_exercise_4_II(get_clues_map)
```

#### Part III (1 point)

Write a function, named `get_number_mines` that computes the number of mines from a map of clues.

**Hint:** this can be written in one line using numpy indexing.

<!-- grade_id: cell-bfa32892688cc39f, points: 0, grade: false, solution: true -->
```python
def get_number_mines(clues_map):
    """
    Computes the number of mines from a map of clues.
    
    Parameters:
    clues_map (np.ndarray): Matrix with map of clues.

    Returns:
    num_mines (int): Number of mines in the board.
    """
    ### BEGIN SOLUTION
    num_mines = len(np.where(clues_map == -1)[0])
    
    return num_mines
    ### END SOLUTION
```

<!-- grade_id: cell-273ffd34483f91fb, points: 1, grade: true, solution: false -->
```python
# This cell has tests to evaluate your answer
# If the cell runs successfully, and no errors are raised, then you get the full score for this exercise
# Otherwise, you are not granted any points


clues_map = np.array([
    [ 1, -1,  1,  0],
    [ 1,  2,  2,  1],
    [ 1,  2, -1,  1],
    [-1,  2,  1,  1],
    [ 1,  1,  0,  0],
])
num_mines = get_number_mines(clues_map)

assert num_mines == 3

from more_tests import test_exercise_4_III
test_exercise_4_III(get_number_mines)
```